## Assignment 1 (Due 5 September 2016 by 11:59pm)

---
## Part a (8 points)
Freedman, SMTP, problems 4.B.7, 4.B.8, 4.B.11, 4.5.3, 4.5.5, 4.5.6,
4.5.10, 4.5.11.

**4.B.7.** 

<!-- grade_id: a_4B7_sol, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**4.B.8.** 

<!-- grade_id: a_4B8_sol, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**4.B.11.** 

<!-- grade_id: a_4B11_sol, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**4.5.3.** 


<!-- grade_id: a_453_sol, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**4.5.5.** 

<!-- grade_id: a_455_sol, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**4.5.6.** 

<!-- grade_id: a_456_sol, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**4.5.10.** 

<!-- grade_id: a_4510_sol, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**4.5.11.** 

<!-- grade_id: a_4511_sol, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

---
## Part b (2 points)
Give an example where floating point addition is not associative.

<!-- grade_id: b_float, points: 0, grade: false, solution: true -->
```python
def example_float_associative():
    """ Assign values to a, b, c, show that float point addition
    is sometimes associative
    a = 
    b = 
    c = 
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    return (a + b) + c == a + (b + c)
```

<!-- grade_id: b_float2, points: 0, grade: false, solution: true -->
```python
def example_float_non_associative():
    """ Assign values to a, b, c, show that float point addition
    can be non-associative
    a = 
    b = 
    c = 
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    return (a + b) + c == a + (b + c)
```

<!-- grade_id: b_float_grade, points: 2, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
assert example_float_associative()
assert not example_float_non_associative()
```

---
## Part c (2 points)
Explain the Kahan Summation Formula in your own words, and explain why it is important.

<!-- grade_id: c_sol, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

---
## Part d (2 points)
Explain why the condition number of a linear system matters.

<!-- grade_id: d_sol, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

---
## Part e (2 points)
Suppose that $A$ is a positive definite square matrix. 

Prove algebraically
that solving the linear system $Ax=y$ by calculating $A^{-1}y$ squares
the condition number.

(Hint: calculating $A^{-1}y$ is a two-step procedure)

<!-- grade_id: e_sol, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

---
## Part f (2 points)
Under what circumstances does it make sense to invert a matrix numerically
rather than solve a linear system by factorization, Gaussian
elimination, back-substitution, or a similar method?

<!-- grade_id: f_sol, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE